<a href="https://colab.research.google.com/github/swaleha/Python/blob/master/Swaleha_quizII_3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **Restart and Run All** (in the menubar in colab, select Runtime$\rightarrow$Restart and Run all) 

Make sure you only fill in any place that says `YOUR CODE HERE` and do not make any other changes to the code. If you add any code cells for your own testing. Please delete them before submission.

---

General Instructions:
- By array, I mean numpy array
- By matrix, I mean a numpy array of 2 dimensions

### 1
Given an array of ints length, return an array with the elements "rotated left" so {1, 2, 3} yields {2, 3, 1}.


rotate_left([1, 2, 3]) → [2, 3, 1] <br>
rotate_left([5, 11, 9]) → [11, 9, 5] <br>
rotate_left([7, 0, 0, 1]) → [0, 0, 1, 7] 

<!-- grade_id: cell-7192257fe41672e9, points: 0, grade: false, solution: true -->
```python
import numpy as np
def rotate_left(num):
  """
  Inputs:
    num : numpy array 
  
  Outputs:
    num_rotated : numpy array
  """
  
  num_rotated = np.roll(num, -1, axis=0)
  
  
  return num_rotated

```

<!-- grade_id: cell-7e19fb61098c249e, points: 5, grade: true, solution: false -->
```python
num = np.array([2, 4, 8, 2, 88, 3, 4])
np.all(rotate_left(num) == np.array([ 4,  8,  2, 88,  3,  4,  2]))
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

### 2
Return the number of times that the string "code" appears anywhere in the given string, except we'll accept any letter for the 'd', so "cope" and "cooe" count.


count_code('aaacodebbb') → 1 <br>
count_code('codexxcode') → 2 <br>
count_code('cozexxcope') → 2

<!-- grade_id: cell-8c217996ba4c3e8f, points: 0, grade: false, solution: true -->
```python
import numpy as np
def count_code(st):
  """
  Inputs:
      st: string
  Outputs:
      cnt: int, number of times 'code' appears in st according to the given specification in the question
  """
  
  sub = ["co"+i+"e" for i in st.lower()]
  cnt_arr = np.char.count(st, sub)
  cnt = np.sum(cnt_arr)
  return cnt
  
```

<!-- grade_id: cell-c976c83c64b32e11, points: 5, grade: true, solution: false -->
```python
st = 'aaacodebbb'
assert count_code(st) == 1
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

### 3
Most frequently occuring element in an array

Given an array, find the most frequent element in it. If there are multiple elements that appear maximum number of times, print any one of them.

Input : arr = [1, 3, 2, 1, 4, 1]
Output : 1

<!-- grade_id: cell-99d7e1295c020be0, points: 0, grade: false, solution: true -->
```python
def most_frequent(arr):
    """
    Inputs:
      arr : numpy array of ints

    Outputs:
      most_freq_el : int, most frequent element in array arr

    """  
    most_freq_el = np.bincount(arr).argmax()
    return most_freq_el


```

<!-- grade_id: cell-0c501da46598a88e, points: 5, grade: true, solution: false -->
```python
arr = [ 1, 5, 2, 1, 3, 2, 1]  
assert most_frequent(arr) == 1
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

### 4
How to find all the local maxima (or peaks) in a 1d array?

Find all the peaks in a 1D numpy array a. Peaks are points surrounded by smaller values on both sides. First and last values of an array cannot be peaks

Input:

arr = [1, 3, 7, 1, 2, 6, 0, 1]

Desired Output: [2, 5]

<!-- grade_id: cell-a3a617509b6c3119, points: 0, grade: false, solution: true -->
```python
def local_maxima(arr):
    """
    Inputs:
        arr: numpy array
    Outputs:
        peak_inds : numpy array of indices of arr where peaks occur
    """
    
    peak_inds = np.where((arr[1:-1] > arr[0:-2]) * (arr[1:-1] > arr[2:]))[0] + 1
    
    return peak_inds
```

<!-- grade_id: cell-8b8d48b1afb4b815, points: 5, grade: true, solution: false -->
```python
assert np.all(local_maxima(np.array([1, 3, 7, 1, 2, 6, 0, 1])) == np.array([2, 5]))
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

### 5
Compute the moving average of n elements of a numpy array <br>
Example:

array:  [8 8 3 7 7 0 4 2 5 2] <br>
moving average for n = 3 :  [ 6.33  6.    5.67  4.67  3.67  2.    3.67  3.  ]

<!-- grade_id: cell-bae1e1468fb9a800, points: 0, grade: false, solution: true -->
```python
def moving_avg(arr, n):
    """
    Inputs:
        arr: numpy array
        n : int, number of elements to consider while computing moving avg
    Outputs:
        moving_avg_arr : numpy array, contains moving avgs
    """
    
    wt = np.repeat(1.0, n) / n  
    moving_avg_arr = np.convolve(arr, wt, 'valid')
    return moving_avg_arr
```

<!-- grade_id: cell-f11b91c4e232dd4a, points: 5, grade: true, solution: false -->
```python
assert np.all(np.isclose(moving_avg(np.array([8, 8, 3, 7, 7, 0, 4, 2, 5, 2]), 3), np.array([6.33333333, 6. , 5.66666667, 4.66666667, 3.66666667, 2. , 3.66666667, 3.]), atol = 0.001))
print('Test passed', '\U0001F44D')

```

Output:
```
Test passed 👍
```

```python

```